Please read the following instructions thoroughly. Neglecting to do so may result in missed points.

### Preamble
**Reminder**: Homeworks are due by 11:59PM ET on Sundays.

Before you turn this problem set in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

### Naming conventions
Be sure the filename of your notebook is in the following form:

    <uni>_<assignment>_<details [optional]>.<extension>
    
For example:

    lr3086_hw01.ipynb
    lr3086_hw01_complete.ipynb
    LR3086_HW01.ipynb
    
To rename a notebook, in the menubar, select File$\rightarrow$Rename. The extension for notebook files, `.ipynb`, will already be appended to the filename, but will be hidden from view within the notebook.
    
This naming format allows for autograding of all assignments. If your files are not named with this format, you should expect a grade of zero for the assignment.

Courseworks may rename your file to something like `lr3086_hw0-1.ipynb` if you resubmit your assignment. This is perfectly fine.

### What Format To Submit In

Most homeworks are in Jupyter notebooks. Once you've finished your homework, unless specified otherwise, please download your work as an `.ipynb` file to your local machine, then upload it to Courseworks when complete (in the menubar, select File$\rightarrow$Download as$\rightarrow$Notebook).

**Failure to submit a Jupyter notebook will result in a grade of zero for the assignment.**

### Grading

Possible points on late assignments are deducted by 50% for each day they are late. For example, if you get 80% of the total possible credits on a homework but hand in that homework a day late, you would get 40%. Assignments two days late get zero points.

Once solutions are posted and graded assignments are handed back, students have 1 week to bring their grading discrepancies to a CA for consideration of possible grading errors.

Because grading is automated, please delete (or comment out) the `raise NotImplmeneted` code before attempting a problem.

Empty un-editable cells in an assignment are there for a reason. They will be filled with tests by the automatic grader. Please do not attempt to remove them.

### Getting Help

Asking for help is a great way to increase your chance of success. However there are some rules. When asking for help (especially from a fellow student), *you can show your helper your code but you can not view theirs*. You work needs to be your own. You can not post screenshots of your current work to Ed Discussions or other tools used for getting help.

If you need to reach out to a CA for help, please do so via Ed Discussions and not via email. Answers given via Ed Discussions will help you as well as other students. Thus, emails will always have a lower priority for response than Ed Discussions questions. If you do email the CA, please make a note of what section you are in. This helps us identify you in Courseworks faster.

Finally, if you do not get a repsonse from a CA within 48 hours, you may email the professor.


---

# Homework 10: SQL

Total questions: 5<br/>
Total points: 8

### FYI

In class, we used paths on your local machine outside of the notebook to persist data, but here, we'll use a special argument that `sqlite3` has which stores the database in-memory -- in otherwords, the database will live only so long as you're running the notebook. Afterwards, it will be thrown away. Usually this argument is useful when experimenting, given that as we discussed, it is persistence that often makes databases useful. We use it here however simply to avoid needing to upload multiple files to Courseworks -- everything you need is here.

If for any reason you need to discard the contents of your database to start over, you may therefore do so by simply restarting the "kernel" of this notebook, which you can do by clicking the restart button in the toolbar (it's 2 buttons to the right of the "Run" button) or in the Kernel menu.

## Question 1

We'll start by creating some tables to store our data -- we'll first have one containing some product information.

Products have:

* a name (`TEXT`)
* a description (`TEXT`)
* a unit cost stored in cents (`INTEGER`)

and of course we also add an id column to identify them.

The SQL to create each table is below. Execute it on the connection provided below to create the table.

```sql
CREATE TABLE IF NOT EXISTS products (
    id INTEGER PRIMARY KEY,
    name TEXT,
    description TEXT,
    unit_cost_cents INTEGER
)
```

[0.5 points]

<!-- grade_id: cell-6cb37374031fab02, points: 0, grade: false, solution: true -->
```python
# Don't not touch these lines!
import sqlite3
connection = sqlite3.connect(":memory:")  # this is our in-memory database, not stored on your hard drive

with connection:
    connection.execute(
        """
        CREATE TABLE IF NOT EXISTS products (
            id INTEGER PRIMARY KEY,
            name TEXT,
            description TEXT,
            unit_cost_cents INTEGER
        )
        """
    )
```

<!-- grade_id: cell-9b9a70f1448562f9, points: 0.5, grade: true, solution: false -->
```python
# CELL INTENTIONALLY LEFT BLANK - DO NOT ALTER OR DELETE
```

## Question 2

Separately, we'd like to track the number of items in stock for each product.

To do so we'll have a `store` and an `inventory` table.

Stores have just an `id` and a `name`. **Create this table** yourself with a similar `CREATE TABLE` statement to the above.

Then, our `inventory` table should combine stores and products, listing how much of each product each store has in stock.

* a `product_id` (`INTEGER`)
* a `store_id` (`INTEGER`)
* a `quantity` (`INTEGER`) in stock

_(Note: in this case below, we actually don't need a separate `id` column because we can use the combination of `store_id` and `product_id` to identify each row. This is known as a_ composite primary key _and is another fundamental relational database concept, but can be ignored if need be for our purposes)._

Here is a statement creating this table:

```sql
CREATE TABLE IF NOT EXISTS inventory (
    product_id INTEGER NOT NULL,
    store_id INTEGER NOT NULL,
    quantity INTEGER,
    FOREIGN KEY (product_id) REFERENCES products(id),
    FOREIGN KEY (store_id) REFERENCES stores(id),
    PRIMARY KEY (product_id, store_id)
)
```
which you should execute as well.

[0.5 points]

<!-- grade_id: cell-9b7c4b5d63e43b02, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
with connection:
    connection.execute(
        """
        CREATE TABLE IF NOT EXISTS stores (
            id INTEGER PRIMARY KEY,
            name TEXT
        )
        """
    )
    connection.execute(
        """
        CREATE TABLE IF NOT EXISTS inventory (
            product_id INTEGER NOT NULL,
            store_id INTEGER NOT NULL,
            quantity INTEGER,
            FOREIGN KEY (product_id) REFERENCES products(id),
            FOREIGN KEY (store_id) REFERENCES stores(id),
            PRIMARY KEY (product_id, store_id)
        )
        """
    )
```

<!-- grade_id: cell-a8bdfa1944f78ff7, points: 0.5, grade: true, solution: false -->
```python
# CELL INTENTIONALLY LEFT BLANK - DO NOT ALTER OR DELETE
```

## Question 3

Now we can insert some stores, products and inventory into our database.

There are 2 stores -- one called `NY` and one called `NJ`. There are 2 products we are concerned with. Their names are `sneakers`, costing `$220` (remember this is dollars!) and boots costing `$350`. Use any description for each that you'd like. `NY` has 4 sneakers in stock and 3 boots. `NJ` has 5 sneakers in stock and no boots.

Insert the above data into the tables you have created.

[4 points]

<!-- grade_id: cell-c9bc54712baa52ca, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE

```

<!-- grade_id: cell-118c7e6641484a93, points: 1, grade: true, solution: false -->
```python
# CELL INTENTIONALLY LEFT BLANK - DO NOT ALTER OR DELETE
```

<!-- grade_id: cell-966bb932a2853739, points: 1, grade: true, solution: false -->
```python
# CELL INTENTIONALLY LEFT BLANK - DO NOT ALTER OR DELETE
```

<!-- grade_id: cell-d3af6e481157efa2, points: 2, grade: true, solution: false -->
```python
# CELL INTENTIONALLY LEFT BLANK - DO NOT ALTER OR DELETE
```

## Question 4

We can of course calculate the total amount of products across all stores.

Write, but do not execute! a `SELECT` query which calculates this sum and returns it as its single returned column. You'll want to use a `SQL` function called `SUM` which functions much like the `COUNT` aggregation function we used in class.

Assign your query to a variable called `query`.

[1.5 point]

<!-- grade_id: cell-edc26d340873c79b, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-782c90de27a48810, points: 1, grade: true, solution: false -->
```python
# CELL INTENTIONALLY LEFT BLANK - DO NOT ALTER OR DELETE
```

<!-- grade_id: cell-bf7419a3e7ed6815, points: 0.5, grade: true, solution: false -->
```python
# CELL INTENTIONALLY LEFT BLANK - DO NOT ALTER OR DELETE
```

## Question 5

What about the total value _in cents_ of all products in stock? I.e. the `sum of (quantity in stock * cost of product across all products)`?

Create a function called `total_current_value` which returns this `int`. You may calculate this entirely in SQL or partially in Python after retrieving the data you need from the database. You must however retrieve some of the data from the database.

[1.5 points]

<!-- grade_id: cell-0bb0b05129715fbd, points: 0, grade: false, solution: true -->
```python
def total_current_value():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-a68838652752ff28, points: 1, grade: true, solution: false -->
```python
### BEGIN TESTS
assert total_current_value() == 303000, total_current_value()
### END TESTS
```

<!-- grade_id: cell-2d9b8bebc6ab900a, points: 0.5, grade: true, solution: false -->
```python
# CELL INTENTIONALLY LEFT BLANK - DO NOT ALTER OR DELETE
```